// variables1.rs
// Make me compile! Execute the command `rustlings hint variables1` if you want a hint :)

// About this `I AM NOT DONE` thing:
// We sometimes encourage you to keep trying things on a given exercise,
// even after you already figured it out. If you got everything working and
// feel ready for the next exercise, remove the `I AM NOT DONE` comment below.

fn main() {
    println!("Hello world");
    let some_string: &str = "manuel";
    println!("My name is {0}, {1} {0}", "Bond", "James");
    println!("Hello, {}", some_string);
    let x = 5;
    println!("x has the value {}", x);
}
